# CP 2025-26: Q2 Lecture 2 - PDE (Part 1)

### General Guidelines

> ⚠️⚠️⚠️ READ CAREFULLY ⚠️⚠️⚠️

- Do not add, delete or create cells, write the answer only in the space marked with the three dots (`...`). Where function skeletons are provided, it is assumed that that function can be called again with different inputs somewhere else. So be careful to write code outside of functions.
  - Function should be ['pure'](https://en.wikipedia.org/wiki/Pure_function), thus no side effects, unless otherwise specified.
- Run the the first cell to import all libraries when opening the notebook before running your own code.
- Read carefully what is required to be printed/returned/plotted in the answer. Please do not output what is not asked for. 
  - If you used the print function for debugging, comment it out ( Ctlr + / ) before submitting
- All plots should have title, xlabel, ylabel, and legend (if there are more than one curve on the plot)
- Use the `help()` function, consult python documentation when using new functions, or do a web search and consult [stackoverflow](https://stackoverflow.com/questions/tagged/python)
- Please read the error messages if you get any, and try to understand what they mean. Debugging code is an essential skill to develop.
- You can use `%debug` to start an IPython console in a cell (or a scratchpad cell!) after an exception has occurred to try to debug.
- You can use `%pdb` to toggle the Python DeBugger (pdb) auto start after an unhandled exception.
- In the assignments you will find some tests put in place, to help you verify your solution. If these fail you are certain you did something wrong, thus look at the hints they provide. But passing these tests does __not__ mean your solution is actually correct.

Make sure you use `python3.12` and the package versions as stated in the provided `requirements.txt`. This file should also be on the course page.

```python
# Importing relevant libraries in the assignment

# This will create static plots (no zooming etc.)
# otherwise try just plain `%matplotlib`, or install a backend such as ipympl or PyQt5 and
# do or `%matplotlib ipympl` `%matplotlib qt`
%matplotlib inline

REPEAT_IMPORTS = True

if REPEAT_IMPORTS or ("IMPORTED_ALL" not in globals()):  # To save you a bit of time

    def print_import_info(package):
        print(
            "Successfully imported %-15s \tVersion: %10s"
            % (package.__name__, package.__version__)
        )

    ### Standard library imports

    import sys

    print("Python version {}".format(sys.version))
    if sys.version_info < (3, 12):
        print(
            "\u001b[31m"  # red
            "\u001b[1m"  # bold
            "WARNING: Use Python 3.12 or newer not to encounter any errors or problems later on. You can chance the the version. This sometime can be done by switching the kernel under the 'Kernel' tab."
            "\u001b[0m"  # reset
        )
    del sys  # Do not need it anymore

    ### Import third party libraries
    # Initialize self assessment helper
    import otter

    grader = otter.Notebook("Assignment_Q2_L2.ipynb")

    import os

    import numpy as np
    import numpy.typing as npt

    print_import_info(np)

    import scipy

    print_import_info(scipy)

    import matplotlib
    import matplotlib.cm as cm
    import matplotlib.pyplot as plt

    print_import_info(matplotlib)

    IMPORTED_ALL = True
    print("Finished importing packages")
else:
    print("Already imported all packages")
```

## Solving transient elliptic PDE - Unsteady 1D heat equation
The 1D unsteady heat equation models the temporal and spatial distribution of temperature in a material, considering how heat diffuses over time along one dimension. In chemical engineering, it is used to analyze processes like the cooling of extruded polymers, temperature regulation in tubular reactors, and thermal management in chemical storage tanks.
\begin{equation}
\frac{dT}{dt} = \alpha\frac{d^2 T}{dx^2}
\tag{Eq. 1}
\end{equation}
where $\alpha$ is the diffusivity of the material in $m^2/s$.

### Numerical solution strategies
To solve the 1D unsteady heat equation numerically using **Finite difference method (FDM)**, we discretize the spatial domain into grid points and the time domain into discrete steps. We can then use time-integrator techniques (Euler method) to update the spatial-discretized temperature in time.
- **Forward Euler (explicit)** method, we update the temperature at each grid point based on the current temperatures, which is straightforward to implement but suffers stability issues.
- **Backward Euler (implicit)** method, on the other hand, involves solving a system of equations at each time step, allowing for larger time steps and greater stability, but at the cost of increased computational and implementation effort.

### Question 1.1: Spatial discretization + Forward (explicit) Euler method
Solve the transient heat equation in 1D using the forward (explicit) Euler method. Temperature is given in Celsius degrees.

\begin{equation}
\frac{dT}{dt} = \alpha\frac{d^2 T}{dx^2}, t \in [0, 0.06] \: , x \in  [0, 3\pi]
\tag{Eq. 4}
\end{equation}

With the following boundary conditions:

$$T\left(x,0\right)  = \sin\left(x\right) $$

$$ T\left(0,t\right)  = 0, \:  t>0$$

$$T\left(L,t\right)  = 0,\: t>0 $$

With:
$$
\alpha = 100 \: m^2/s
$$


#### Instructions

1. **Discretization File**:
   - Create a new Python file named **`discretization.py`**.
   - Define the following functions in `discretization.py` to set up the discretized system:
     - **`define_A_DF(nx: int, fo: float) -> np.ndarray`**: Returns the matrix `A_DF` required for discretizing the internal spatial points.
     - **`define_B_DF(nx: int, lbc: float, rbc: float) -> np.ndarray`**: Returns the vector `b_DF` that incorporates the Dirichlet boundary conditions at the left and right boundaries of the domain.
   - Both functions should be structured to handle only the internal points, making the dimensionality of `A_DF` and `b_DF` equal to `nx - 2`.

2. **Import Functions into the Notebook**:
   - In the current Jupyter notebook, import the functions `define_A_DF` and `define_B_DF` from the file `discretization.py`. These will be used as part of the numerical solution process for the heat equation.

3. **Implement the Euler Forward Integration**:
   - Define a function in your notebook called `euler_explicit` with the following inputs and outputs:

     **Arguments**:
     - `x_mesh` (`npt.NDArray`): A 1D array of grid points representing the spatial domain.
     - `t_mesh` (`npt.NDArray`): A 1D array of grid points representing the time domain.
     - `alpha` (`float`): The thermal diffusivity constant in the heat equation.

     **Returns**:
     - `npt.NDArray`: A 2D array where each entry represents the temperature at a given time and spatial point, with dimensions `[len(t_mesh), len(x_mesh)]`.

   - **Function Overview**:
     - The function `euler_explicit` should initialize the temperature profile across the spatial domain.
     - For each time step in `t_mesh`, update the temperature profile using the explicit Euler method, applying the matrix `A_DF` and vector `b_DF` derived from `discretization.py`.

```python
# Import the necessary functions
...
```

```python
# Define the function `euler_explicit`
...
```

```python
grader.check("q1_1_1")
```

**Now that you Forward explicit euler method is correct, try to solve the problem above**

```python
# Input variables and solution variable
x_mesh = np.linspace(0, 3 * np.pi, 501)
t_mesh = np.linspace(0, 0.06, 501)
alpha = 100.0
results_forward = euler_explicit(x_mesh, t_mesh, alpha)

if np.any(np.isnan(results_forward)):
    print("The method does not converge!")
```

### Reflect on the results
1. Describe what you observe in the results and provide an intuition why this is happening. Then proceed with the assignment to mitigate the issue.

_Type your answer here, replacing this text._

How many time discretization points would be needed (at least) to have the Explicit Euler method converging (assuming 501 spatial discretization points and the same spatial and temporal domain as above)?

```python
...
n_time_points_min = ...
```

```python
grader.check("q1_1_2")
```

### Question 1.2: Spatial discretization + Backward (implicit) Euler method
#### Instructions

1. **Update `discretization.py` for the Backward Euler Method**:
   - In the existing file **`discretization.py`**, add two new functions to create the matrix and vector needed for the backward Euler method:
     - **`define_A_DB(nx: int, fo: float) -> np.ndarray`**: Returns the matrix `A_DB` for discretizing the internal spatial points using the backward Euler method.
     - **`define_B_DB(b_DF: np.ndarray, yk: np.ndarray, fo: float) -> np.ndarray`**: Returns the vector `b_DB` that incorporates the Dirichlet boundary conditions and previous time step values.

2. **Import Functions into the Notebook**:
   - In the current Jupyter notebook, import the newly added functions `define_A_DB` and `define_B_DB` from `discretization.py` along with the functions for the forward Euler method (`define_A_DF` and `define_B_DF`).

3. **Implement the Euler Backward Integration**:
   - Define a function in your notebook called `euler_implicit` with the following inputs and outputs:

     **Arguments**:
     - `x_mesh` (`npt.NDArray`): A 1D array of grid points representing the spatial domain.
     - `t_mesh` (`npt.NDArray`): A 1D array of grid points representing the time domain.
     - `alpha` (`float`): The thermal diffusivity constant in the heat equation.

     **Returns**:
     - `npt.NDArray`: A 2D array where each entry represents the temperature at a given time and spatial point, with dimensions `[len(t_mesh), len(x_mesh)]`.

   - **Function Overview**:
     - The function `euler_implicit` should initialize the temperature profile across the spatial domain.
     - For each time step in `t_mesh`, update the temperature profile using the implicit Euler method, applying the matrix `A_DB` and vector `b_DB` derived from `discretization.py`.
     - Solve the linear system that arises from the implicit scheme at each time step to advance the temperature profile in time.

*Don't Repeat Yourself (DRY) principle:  Can you re-use some of the function implemented for the Forward Euler case?*

```python
# Import the necessary functions
...
```

```python
# Define the function `euler_implicit`
...
```

```python
# Input variables and solution variable
x_mesh = np.linspace(0, 3 * np.pi, 5)
t_mesh = np.linspace(0, 0.06, 5)
alpha = 100.0
results_backward = euler_implicit(x_mesh, t_mesh, alpha)

# Plotting and plot handles
fig, ax = plt.subplots(subplot_kw={"projection": "3d"})
XX, TT = np.meshgrid(x_mesh, t_mesh)
surf = ax.plot_surface(XX, TT, results_backward, cmap=cm.coolwarm)
MX = int(x_mesh[-1] / np.pi)
plt.xticks(np.arange(0, MX + 1) * np.pi, labels=list(map(str, np.arange(0, MX + 1))))
plt.xlabel(r"x$ [m]$")
ax.set_zlabel(r"$T [\degree C]$")
plt.ylabel(r"$t [s]$")
plt.suptitle("Solution to the transient heat equation using Euler implicit method")
plot_handles_12 = (fig, [ax])
```

```python
grader.check("q1_2")
```

### Reflect on the results
1. The Backward Euler method requires solving a linear system at each time step. Explain why this is computationally more expensive than Forward Euler, and suggest strategies to reduce the cost.

_Type your answer here, replacing this text._

### Part 2: Implementing Backward Euler with Sparse Matrix Definition

**Overview**

In this exercise, you will use a sparse matrix representation to implement the backward Euler method for the heat equation. Sparse matrices store only non-zero entries, which reduces memory usage and computation time, making this approach highly efficient for large spatial grids and large-scale linear systems.

#### Instructions

1. **Add Sparse Matrix Function to `discretization.py`**:
   - Add a new function in your `discretization.py` file with the following header and details:

     **Function Name**: `define_A_DB_sparse`

     **Arguments**:
     - `nx` (`int`): Number of spatial discretization points.
     - `fo` (`float`): Fourier number (dimensionless), which controls time-step stability.

     **Returns**:
     - `scipy.sparse.csc_matrix`: A sparse matrix of shape `(nx-2, nx-2)` for the Dirichlet Backward Euler method, created using `scipy.sparse.spdiags`. This matrix represents the discretization for the implicit Euler method in sparse format.

     - **Description**:
       - This function defines the sparse matrix `A_DB` using `scipy.sparse.spdiags`, leveraging sparse storage to reduce memory and computation requirements.
       - For more information on `spdiags`, refer to the [documentation](https://docs.scipy.org/doc/scipy/reference/generated/scipy.sparse.spdiags.html).

2. **Implement the Sparse Backward Euler Integration in the Notebook**:
   - After importing `define_A_DB_sparse` from `discretization.py`, implement a function in your notebook with the following header and details:

     **Function Name**: `euler_implicit_sparse`

     **Arguments**:
     - `x_mesh` (`npt.NDArray`): A 1D array of grid points representing the spatial domain.
     - `t_mesh` (`npt.NDArray`): A 1D array of grid points representing the time domain.
     - `alpha` (`float`): The thermal diffusivity constant in the heat equation.

     **Returns**:
     - `npt.NDArray`: A 2D array with each entry representing the temperature at a given time and spatial point, with dimensions `[len(t_mesh), len(x_mesh)]`.

   - **Function Overview**:
     - The function `euler_implicit_sparse` initializes the temperature profile across the spatial domain.
     - For each time step in `t_mesh`, it updates the temperature profile by solving a linear system with the sparse matrix `A_DB_sparse` and a vector `b_DB`, which includes contributions from boundary conditions and the previous time step.
     - Use the `scipy.sparse.linalg.spsolve` function to solve the sparse linear system efficiently. For details, refer to the [spsolve documentation](https://docs.scipy.org/doc/scipy/reference/generated/scipy.sparse.linalg.spsolve.html).

---

**Special Note**: Using sparse matrices with efficient solvers optimizes both memory usage and computational efficiency, making it ideal for large-scale scientific problems. This exercise will help you apply these principles to solve the heat equation with high performance and scalability.


```python
# Import the necessary functions
...
```

```python
# Define the function `euler_implicit_sparse`
...
```

```python
# Input variables and solution variable
x_mesh = np.linspace(0, 3 * np.pi, 501)
t_mesh = np.linspace(0, 0.06, 501)
alpha = 100.0
results_backward = euler_implicit_sparse(x_mesh, t_mesh, alpha)

# Plotting and plot handles
fig, ax = plt.subplots(subplot_kw={"projection": "3d"})
# BEGIN SOLUTION
XX, TT = np.meshgrid(x_mesh, t_mesh)
surf = ax.plot_surface(XX, TT, results_backward, cmap=cm.coolwarm)
MX = int(x_mesh[-1] / np.pi)
plt.xticks(np.arange(0, MX + 1) * np.pi, labels=list(map(str, np.arange(0, MX + 1))))
plt.xlabel(r"x$ [m]$")
ax.set_zlabel(r"$T [K]$")
plt.ylabel(r"$t [s]$")
plt.suptitle("Solution to the transient heat equation using Euler implicit method")
plot_handles_12 = (fig, [ax])
```

**Check yourself: Do you get the same solution as before? Is the computation running faster than before?**

*Insert your answer here*

Tests for q1_1_1 (points: 0)
```
>>> assert os.path.exists('discretization.py'), 'File discretization.py does not exist. Are you sure that you placed in the same directory as this notebook?'
```
```
>>> from discretization import define_A_DF, define_B_DF
>>> assert callable(define_A_DF), 'Function define_A_DF is not defined in discretization.py.'
```
```
>>> from discretization import define_A_DF, define_B_DF
>>> assert callable(define_B_DF), 'Function define_B_DF is not defined in discretization.py.'
```
```
>>> from discretization import define_A_DF, define_B_DF
>>> nx = 10
>>> fo = 0.1
>>> lbc = 1.0
>>> rbc = 0.0
>>> expected_shape_A = (nx - 2, nx - 2)
>>> A_DF = define_A_DF(nx, fo)
>>> assert A_DF.shape == expected_shape_A, f'Matrix A_DF has shape {A_DF.shape}, expected {expected_shape_A}.'
```
```
>>> from discretization import define_A_DF, define_B_DF
>>> nx = 10
>>> fo = 0.1
>>> lbc = 1.0
>>> rbc = 0.0
>>> b_DF = define_B_DF(nx, lbc, rbc)
>>> expected_shape_b = (nx - 2,)
>>> assert b_DF.shape == expected_shape_b, f'Vector b_DF has shape {b_DF.shape}, expected {expected_shape_b}.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_explicit(x_mesh, t_mesh, alpha)
>>> expected_shape = (len(t_mesh), len(x_mesh))
>>> assert results_forward.shape == expected_shape, f'Expected shape {expected_shape}, but got {results_forward.shape}'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_explicit(x_mesh, t_mesh, alpha)
>>> initial_condition = np.sin(x_mesh)
>>> assert np.allclose(results_forward[0, :], initial_condition), 'Initial condition u(x, 0) = sin(x) was not applied correctly.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_explicit(x_mesh, t_mesh, alpha)
>>> assert np.all(results_forward[:, 0] == 0), 'Left boundary condition u(0, t) = 0 was not maintained.'
>>> assert np.all(results_forward[:, -1] == 0), 'Right boundary condition u(L, t) = 0 was not maintained.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_explicit(x_mesh, t_mesh, alpha)
>>> assert np.all(np.isfinite(results_forward)), 'The results array contains non-finite values (NaN or Inf).'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_explicit(x_mesh, t_mesh, alpha)
>>> def __euler_explicit(x_mesh, t_mesh, alpha):
...     Nx = len(x_mesh)
...     Nt = len(t_mesh)
...     results = np.empty((Nt, Nx))
...     dx = x_mesh[1] - x_mesh[0]
...     dt = t_mesh[1] - t_mesh[0]
...     Fo = alpha * dt / dx ** 2
...     A_DF = np.diag(np.full(Nx - 2, 1 / Fo - 2))
...     A_DF += np.diag(np.ones(Nx - 3), k=1)
...     A_DF += np.diag(np.ones(Nx - 3), k=-1)
...     b_DF = np.zeros(Nx - 2)
...     b_DF[0] = 0
...     b_DF[-1] = 0
...     results[0, :] = np.sin(x_mesh)
...     results[:, 0] = 0
...     results[:, -1] = 0
...     for time_idx, t in enumerate(t_mesh[1:], 1):
...         results[time_idx, 1:-1] = Fo * np.dot(A_DF, results[time_idx - 1, 1:-1]) + Fo * b_DF
...     return results
>>> __results_forward = __euler_explicit(x_mesh, t_mesh, alpha)
>>> assert np.allclose(results_forward, __results_forward), 'Your explicit Euler method does not give the correct result.'
```


Tests for q1_1_2 (points: 0)
```
>>> assert isinstance(n_time_points_min, int), 'The minimal number of time discretization points must be an integer.'
```
```
>>> __n_time_points_min = 33774
>>> assert __n_time_points_min == n_time_points_min, 'Your estimation of minimal number of time discretization points is incorrect.'
```


Tests for q1_2 (points: 0)
```
>>> from discretization import define_A_DB, define_B_DB
>>> assert callable(define_A_DB), 'Function define_A_DB is not defined in discretization.py.'
```
```
>>> from discretization import define_A_DB, define_B_DB
>>> assert callable(define_B_DB), 'Function define_B_DB is not defined in discretization.py.'
```
```
>>> from discretization import define_A_DB, define_B_DB
>>> nx = 10
>>> fo = 0.1
>>> lbc = 1.0
>>> rbc = 0.0
>>> expected_shape_A = (nx - 2, nx - 2)
>>> A_DB = define_A_DB(nx, fo)
>>> assert A_DB.shape == expected_shape_A, f'Matrix A_DB has shape {A_DF.shape}, expected {expected_shape_A}.'
```
```
>>> from discretization import define_A_DB, define_B_DB, define_B_DF
>>> nx = 10
>>> fo = 0.1
>>> lbc = 1.0
>>> rbc = 0.0
>>> b_DF = define_B_DF(nx, lbc, rbc)
>>> b_DB = define_B_DB(b_DF, np.zeros(nx - 2), fo)
>>> expected_shape_b = (nx - 2,)
>>> assert b_DB.shape == expected_shape_b, f'Vector b_DB has shape {b_DF.shape}, expected {expected_shape_b}.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_implicit(x_mesh, t_mesh, alpha)
>>> expected_shape = (len(t_mesh), len(x_mesh))
>>> assert results_forward.shape == expected_shape, f'Expected shape {expected_shape}, but got {results_forward.shape}'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_implicit(x_mesh, t_mesh, alpha)
>>> initial_condition = np.sin(x_mesh)
>>> assert np.allclose(results_forward[0, :], initial_condition), 'Initial condition u(x, 0) = sin(x) was not applied correctly.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_implicit(x_mesh, t_mesh, alpha)
>>> assert np.all(results_forward[:, 0] == 0), 'Left boundary condition u(0, t) = 0 was not maintained.'
>>> assert np.all(results_forward[:, -1] == 0), 'Right boundary condition u(L, t) = 0 was not maintained.'
```
```
>>> x_mesh = np.linspace(0, 3 * np.pi, 51)
>>> t_mesh = np.linspace(0, 1, 51)
>>> alpha = 1e-05
>>> results_forward = euler_implicit(x_mesh, t_mesh, alpha)
>>> assert np.all(np.isfinite(results_forward)), 'The results array contains non-finite values (NaN or Inf).'
```
```
>>> __x_mesh = np.linspace(0, 2 * np.pi, 101)
>>> __t_mesh = np.linspace(0, 0.1, 101)
>>> __alpha = 150.0
>>> assert isinstance(euler_implicit(x_mesh, t_mesh, alpha), np.ndarray), "Your function 'euler_implicit' does not return an array."
>>> assert len(euler_implicit(__x_mesh, __t_mesh, alpha)[0]) == 101, "Your function 'euler_implicit' does not change with the input."
>>> assert len(euler_implicit(__x_mesh, __t_mesh, alpha)[:, 0]) == 101, "Your function 'euler_implicit' does not change with the input."
>>> assert np.allclose(euler_implicit(__x_mesh, __t_mesh, alpha), euler_implicit(__x_mesh, __t_mesh, __alpha)) is False, "Your function 'euler_implicit' does not change with the input."
```
