# CP 2025-26: Assignment 7 - Boundary value problems (Shooting method)

### General Guidelines

> ⚠️⚠️⚠️ READ CAREFULLY ⚠️⚠️⚠️

- Do not add, delete or create cells, write the answer only in the space marked with the three dots (`...`). Where function skeletons are provided, it is assumed that that function can be called again with different inputs somewhere else. So be careful to write code outside of functions.
  - Function should be ['pure'](https://en.wikipedia.org/wiki/Pure_function), thus no side effects, unless otherwise specified.
- Run the the first cell to import all libraries when opening the notebook before running your own code.
- Read carefully what is required to be printed/returned/plotted in the answer. Please do not output what is not asked for. 
  - If you used the print function for debugging, comment it out ( Ctlr + / ) before submitting
- All plots should have title, xlabel, ylabel, and legend (if there are more than one curve on the plot)
- Use the `help()` function, consult python documentation when using new functions, or do a web search and consult [stackoverflow](https://stackoverflow.com/questions/tagged/python)
- Please read the error messages if you get any, and try to understand what they mean. Debugging code is an essential skill to develop.
- You can use `%debug` to start an IPython console in a cell (or a scratchpad cell!) after an exception has occurred to try to debug.
- You can use `%pdb` to toggle the Python DeBugger (pdb) auto start after an unhandled exception.
- In the assignments you will find some tests put in place, to help you verify your solution. If these fail you are certain you did something wrong, thus look at the hints they provide. But passing these tests does __not__ mean your solution is actually correct.

Make sure you use `python3.12` and the package versions as stated in the provided `requirements.txt`. This file should also be on the course page.

```python
# Importing relevant libraries in the assignment

# This will create static plots (no zooming etc.)
# otherwise try just plain `%matplotlib`, or install a backend such as ipympl or PyQt5 and
# do or `%matplotlib ipympl` `%matplotlib qt`
%matplotlib inline

REPEAT_IMPORTS = True

if REPEAT_IMPORTS or ("IMPORTED_ALL" not in globals()):  # To save you a bit of time

    def print_import_info(package):
        print(
            "Successfully imported %-15s \tVersion: %10s"
            % (package.__name__, package.__version__)
        )

    ### Standard library imports

    import sys

    print("Python version {}".format(sys.version))
    if sys.version_info < (3, 12):
        print(
            "\u001b[31m"  # red
            "\u001b[1m"  # bold
            "WARNING: Use Python 3.12 or newer to not encounter any errors or "
            "problems later on.\n"
            "\u001b[0m"  # reset
        )
    del sys  # Do not need it anymore

    import typing
    from typing import Callable, List, Tuple, Union

    ### Import third party libraries
    # Initialize self assessment helper
    import otter

    grader = otter.Notebook("Assignment_Q1_L7.ipynb")

    import numpy as np
    import numpy.typing as npt

    print_import_info(np)

    import scipy
    import scipy.integrate

    print_import_info(scipy)

    import matplotlib
    import matplotlib.pyplot as plt

    print_import_info(matplotlib)

    IMPORTED_ALL = True
    print("Finished importing packages")
else:
    print("Already imported all packages")
```

## Question 1: Convection-diffusion equation

We aim to solve the **steady one-dimensional convection–diffusion** equation on the interval $x \in [0,1]$. Our objective is to numerically approximate the function $u(x)$, which satisfies:

\begin{equation}
\tag{Eq. 1}
 -\varepsilon\cfrac{\text{d}^2u}{\text{d}x^2}+\cfrac{\text{d}u}{\text{d}x} = 0
\end{equation}

The parameter $\varepsilon>0$ controls the relative strength of the diffusion and convection terms.
In computational transport problems, $\varepsilon$ often appears as the inverse of a Péclet number (or sometimes a Reynolds number) and strongly influences the character of the solution.

In this assignment, set $\varepsilon=0.1$

### Boundary Conditions

The boundary conditions are:

\begin{equation}
\tag{Eq. 2}
 u(x=0)=1, \quad u(x=1)=0
\end{equation}

That is, a **non-homogeneous Dirichlet** condition at $x=0$ and a **homogeneous Dirichlet** condition at $x=1$.

### Analytical solution

The differential equation $(\text{Eq.1})$ with the boundary conditions $(\text{Eq.2})$ can be solved analytically by assuming a solution of the form $u(x) = \exp(r x)$. Substituting into $(\text{Eq.1})$ yields the characteristic equation: $-\varepsilon r^2 +r = 0$ with roots $r_1=0$ and $r_2=1/\varepsilon$.

Hence, the general solution is: $u(x) = C_{1} \exp(r_{1} x)+C_{2} \exp(r_{2} x)$. Applying the boundary condition give the final form of the analytical solution:

\begin{equation}
\tag{Eq. 3}
 u(x)=\frac{\exp (x/\varepsilon)-\exp (1/\varepsilon)}{1-\exp (1/\varepsilon)} 
\end{equation}

### Question 1.1: Shooting method

Solve the convection-diffusion equation  ($\text{Eq.1}$) using the shooting method.

Your code should:
- Define the parameter $\varepsilon$ as `eps=0.1`. **See note below**.
- Create a uniform grid `x` using **101 equidistant grid points** over the interval $x \in [0,1]$.
- Define a function `ode_system(x,y)` that defines the second-order ODE in ($\text{Eq.1}$) as a system of first-order ODEs. Include type hints and a docstring.
- Set boundary conditions $u(x=0)=1$ and $u(x=1)=0$ using variables `u0` and `u1`.
- Implement the shooting method using the **secant method** as a root-finding technique:
  - A tolerance `tol=1e-5`.
  - A maximum iteration count `max_iter=100`.
- Store the converged numerical solution in the variable `shooting_solution`. 
- Compute and store the analytical solution in the variable `analytical_solution`. 
- Make a plot of the numerical solution against the analytical solution. Include a title, legend and axis labels.

**NOTE**: In this assignment, you are asked to define `eps` as a global parameter so that it does not need to be passed explicitly to the functions that define the system of ODEs. Normally, using global variables is discouraged because it can make code harder to maintain and debug. However, in this particular case, defining `eps` globally simplifies the implementation with `solve_bvp()` later on in the assignment, since passing additional parameters to this solver function is less straightforward than with `solve_ivp()`. 

```python
# Parameters
eps = ...
x = ...

# ODE system
...

# Boundary conditions for u at x=0 and x=1
u0 = ...
u1  = ...

# Shooting method implementation
tol = ...
max_iter = ...

...

# Extract shooting solution
shooting_solution = ...

# Analytical solution
analytical_solution = ...

# Plot results
...
```

```python
grader.check("q1_1")
```

### Question 1.2: solve_bvp()

Solve the convection-diffusion equation ($\text{Eq.1}$) using the in-built function in scipy [`solve_bvp()`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.integrate.solve_bvp.html).

Your code should:
- Reuse the parameter $\varepsilon$ `eps`, grid `x` and the ODE system function `ode_system(x,y)` from the previous question.
- Define a boundary condition function `bc(u0,u1)` that returns the residuals of the boundary conditions. Remember to define type hints and docstrings. Here `u0` defines the function value at $x=0$ and `u1` defines the function value at $x=1$.
- Solve the system using `solve_bvp()` and store the resulting solution in the variable `bvp_solution`. Remember to define an initial guess for the `solve_bvp()` function.
- Plot your `bvp_solution` together with the shooting method solution `shooting_solution` obtained in the previous question.  Include a title, legend and axis labels.

```python
# Define boundary conditions function for solve_bvp
...

# Solve using solve_bvp
...

# Extract bvp solution
bvp_solution = ...

# Plot results
...
```

```python
grader.check("q1_2")
```

### Reflect on the results

The shooting method transforms a boundary value problem (BVP) into an initial value problem (IVP). For small $\epsilon$, the convection–diffusion equation develops a sharp boundary layer near $x = 1$.
1. Explain why this makes the shooting method highly sensitive to the initial slope at $x = 0$, and predict what happens if your guess is even slightly wrong.
2. Would solve_bvp() show the same kind of sensitivity to the initial slope guess? 

Justify your answer in 4–6 sentences in the text box below.



_Type your answer here, replacing this text._

## Question 2

We now solve the same steady one-dimensional convection–diffusion system defined in Question 1:

\begin{equation}
\tag{Eq. 1}
 -\varepsilon\cfrac{\text{d}^2u}{\text{d}x^2}+\cfrac{\text{d}u}{\text{d}x} = 0,
\end{equation}

but with a **Neumann** boundary condition at $x=0$. The new set of boundary conditions are thus defined as: 

\begin{equation}
\tag{Eq. 4}
\frac{du}{dx}\bigg\vert_{x=0}=-0.001, \quad u(x=1)=0.
\end{equation}

### Analytical Solution

The analytical solution for this new set of boundary conditions becomes:

\begin{equation}
\tag{Eq. 3}
 u(x)=-0.001\varepsilon\left(\exp(x/\varepsilon)-\exp(1/\varepsilon)\right)
\end{equation}

### Question 2.1: Shooting Method 

Solve $\text{Eq.1}$, given the new boundary conditions from $\text{Eq.4}$, using the shooting method. 

Your code should:
- Reuse the parameter $\varepsilon$ `eps`, grid `x` and the ODE system function `ode_system(x,y)` from the previous question.
- Impose a Neumann boundary condition $\frac{du}{dx}(x=0)=-0.001$ and a Dirichlet boundary condition $u(x=1)=0$. Store these using variables `du_dx_0` and `u1`.
- Implement the shooting method using the **secant method** as a root-finding technique and using the same tolerance `tol=1e-5` and maximum iteration count `max_iter=100` as in Question 1.
- Store the converged numerical solution in the variable `shooting_solution2`. 
- Compute and store the analytical solution in the variable `analytical_solution2`. 
- Make a plot of the numerical solution against the analytical solution. Include a title, legend and axis labels.

```python
# Boundary conditions
du_dx_0 = ...
u1 = ...

# Shooting method implementation for Neumann BC
...

# Extract shooting solution
shooting_solution2 = ...

# Analytical solution
analytical_solution2 = ...

# Plot results
...
```

```python
grader.check("q2_1")
```

### Question 2.2: solve_bvp()

Solve $\text{Eq.1}$, given the new boundary conditions from $\text{Eq.4}$, using the in-built function in scipy [`solve_bvp()`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.integrate.solve_bvp.html).

Your code should:
- Reuse the parameter $\varepsilon$ `eps`, grid `x` and the ODE system function `ode_system(x,y)` from the previous question.
- Define a boundary condition function `bc2(u0,u1)` that returns the residuals of the boundary conditions. Remember to define type hints and docstrings. Here `u0` defines the function value at $x=0$ and `u1` defines the function value at $x=1$.
- Solve the system using `solve_bvp()` and store the resulting solution in the variable `bvp_solution2`. Remember to define an initial guess for the `solve_bvp()` function.
- Plot your `bvp_solution2` together with the shooting method solution `shooting_solution2` obtained in the previous question.  Include a title, legend and axis labels.

```python
# Define boundary conditions function for solve_bvp with Neumann BC
...

# Solve using solve_bvp
...

# Extract bvp solution
bvp_solution2 = ...

# Plot results
...
```

```python
grader.check("q2_2")
```

## Bonus Question 3

Solve $\text{Eq.1}$, using the same boundary conditions from Question 1 given in $\text{Eq.2}$, with the finite difference method. As in Assignment 6, construct the coefficient matrix $A$ so that you solve only for the unknown points of $u(x)$.

Your code should:
- Reuse the parameter $\varepsilon$ `eps` and grid `x` from the previous questions.
- Store the numerical solution $u(x)$ obtained using the finite difference method in the variable `u_fd`. 
- Plot your numerical solution together with the analytical solution computed earlier (`analytical_solution`).  Include a title, legend and axis labels.

**NOTE**: For this question, only grader checks for the final solution `u_fd` are implemented.

```python
# Finite Difference Method Implementation
...

# Plot results
...
```

```python
grader.check("q3_1")
```

```python

```

Tests for q1_1 (points: None)
```
>>> assert np.isclose(eps, 0.1), f'`eps` should be 0.1, got {eps}.'
>>> assert isinstance(x, np.ndarray), '`x` should be a NumPy array.'
>>> assert len(x) == 101, f'`x` should contain 101 points, got {len(x)}.'
>>> assert np.isclose(x[0], 0.0) and np.isclose(x[-1], 1.0), 'Grid `x` must span from 0 to 1.'
>>> assert np.allclose(np.diff(x), np.diff(x)[0]), 'Grid `x` must be uniformly spaced.'
```
```
>>> test_x = 0.5
>>> test_y = np.array([1.0, 2.0])
>>> expected_output = np.array([2.0, 2.0 / eps])
>>> result = ode_system(test_x, test_y)
>>> assert isinstance(result, np.ndarray), '`ode_system` must return a NumPy array.'
>>> assert np.allclose(result, expected_output), f'`ode_system` output incorrect!'
```
```
>>> assert np.isclose(u0, 1.0), f'`u0` should be 1.0, got {u0}.'
>>> assert np.isclose(u1, 0.0), f'`u1` should be 0.0, got {u1}.'
```
```
>>> assert isinstance(shooting_solution, np.ndarray), '`shooting_solution` must be a NumPy array.'
>>> assert shooting_solution.shape == x.shape, '`shooting_solution` must have the same length as `x`.'
>>> __shooting_solution = np.array([1.0, 0.999995215, 0.999989927, 0.999984082, 0.999977624, 0.999970485, 0.999962593, 0.999953864, 0.999944222, 0.999933578, 0.999921827, 0.999908848, 0.999894507, 0.999878651, 0.999861116, 0.999841719, 0.999820263, 0.999796535, 0.999770308, 0.999741338, 0.999709319, 0.999673894, 0.999634779, 0.999591627, 0.999544021, 0.999491479, 0.999433448, 0.999369308, 0.999298374, 0.999219888, 0.999133028, 0.999036902, 0.998930551, 0.998812948, 0.998682997, 0.998539536, 0.99838095, 0.998205532, 0.998011886, 0.997798283, 0.997562657, 0.997302605, 0.99701539, 0.996697934, 0.996346827, 0.995958321, 0.995528331, 0.995052436, 0.994525879, 0.993943567, 0.993300068, 0.992589617, 0.99180466, 0.990936164, 0.989977243, 0.988919433, 0.987752586, 0.986464869, 0.985042764, 0.983471067, 0.981732892, 0.979809666, 0.97768113, 0.975325343, 0.972718678, 0.969835821, 0.966649776, 0.963131861, 0.959246155, 0.954946385, 0.950198192, 0.944960017, 0.939181895, 0.932805462, 0.925763949, 0.917982185, 0.909376598, 0.89985521, 0.889317644, 0.877655119, 0.864750451, 0.850478054, 0.834703939, 0.817285715, 0.798048851, 0.776761776, 0.753252764, 0.727316386, 0.69870561, 0.667131794, 0.63226469, 0.593732447, 0.551121603, 0.503977093, 0.451802244, 0.394058777, 0.330166807, 0.259504843, 0.181409787, 0.095176933, 4.8507478e-07])
>>> assert np.allclose(shooting_solution, __shooting_solution, rtol=0.01, atol=0.01), ' `shooting_solution` does not match expected values.'
```
```
>>> assert isinstance(analytical_solution, np.ndarray), '`analytical_solution` must be a NumPy array.'
>>> assert analytical_solution.shape == x.shape, '`analytical_solution` must have the same length as `x`.'
>>> __analytical_solution = arr = np.array([1.0, 0.99999523, 0.99998995, 0.99998412, 0.99997767, 0.99997055, 0.99996267, 0.99995397, 0.99994436, 0.99993373, 0.99992199, 0.99990901, 0.99989466, 0.99987881, 0.99986129, 0.99984192, 0.99982052, 0.99979687, 0.99977074, 0.99974185, 0.99970992, 0.99967464, 0.99963565, 0.99959255, 0.99954493, 0.99949229, 0.99943412, 0.99936983, 0.99929878, 0.99922026, 0.99913348, 0.99903757, 0.99893158, 0.99881443, 0.99868497, 0.99854189, 0.99838377, 0.99820901, 0.99801588, 0.99780243, 0.99756654, 0.99730583, 0.99701771, 0.99669928, 0.99634737, 0.99595845, 0.99552862, 0.99505358, 0.99452859, 0.99394838, 0.99330715, 0.99259848, 0.99181528, 0.99094971, 0.98999311, 0.9889359, 0.9877675, 0.98647623, 0.98504914, 0.98347197, 0.98172893, 0.97980257, 0.97767361, 0.97532075, 0.97272044, 0.96984665, 0.96667062, 0.96316056, 0.95928135, 0.95499415, 0.95025607, 0.94501968, 0.93923258, 0.93283684, 0.92576845, 0.91795668, 0.90932333, 0.89978201, 0.88923721, 0.87758341, 0.86470397, 0.85046999, 0.83473901, 0.81735358, 0.79813972, 0.77690511, 0.75343724, 0.72750124, 0.69883752, 0.66715921, 0.63214926, 0.59345728, 0.55069604, 0.50343755, 0.45120885, 0.3934872, 0.32969492, 0.25919355, 0.18127748, 0.0951669, -0.0])
>>> assert np.allclose(analytical_solution, __analytical_solution, rtol=0.01, atol=0.01), ' `analytical_solution` solution does not match expected values.'
```


Tests for q1_2 (points: None)
```
>>> u_left_test = np.array([1.0, 0.0])
>>> u_right_test = np.array([0.0, 0.0])
>>> result = bc(u_left_test, u_right_test)
>>> assert isinstance(result, np.ndarray), '`bc` must return a NumPy array.'
>>> assert result.shape == (2,), f'`bc` must return a 2-element array, got shape {result.shape}.'
>>> expected = np.array([u_left_test[0] - 1, u_right_test[0]])
>>> assert np.allclose(result, expected), f'`bc` output is incorrect.'
```
```
>>> assert isinstance(bvp_solution, np.ndarray), '`bvp_solution` must be a NumPy array.'
>>> assert bvp_solution.shape == x.shape, '`bvp_solution` must have the same length as `x`.'
>>> __bvp_solution = np.array([1.0, 0.999995225, 0.999989948, 0.999984116, 0.99997767, 0.999970547, 0.999962674, 0.999953974, 0.999944358, 0.999933731, 0.999921986, 0.999909007, 0.999894662, 0.999878808, 0.999861288, 0.999841924, 0.999820524, 0.999796874, 0.999770736, 0.999741849, 0.999709924, 0.999674641, 0.999635648, 0.999592554, 0.999544927, 0.999492292, 0.999434121, 0.999369832, 0.999298782, 0.999220259, 0.999133478, 0.99903757, 0.998931575, 0.998814433, 0.998684971, 0.998541893, 0.998383768, 0.998209012, 0.998015878, 0.997802431, 0.997566535, 0.997305831, 0.997017707, 0.996699282, 0.996347368, 0.995958442, 0.995528613, 0.995053578, 0.994528583, 0.993948374, 0.993307144, 0.992598476, 0.991815276, 0.990949706, 0.989993104, 0.988935894, 0.987767497, 0.986476219, 0.985049136, 0.983471965, 0.981728921, 0.979802561, 0.977673603, 0.97532074, 0.972720425, 0.969846633, 0.966670601, 0.963160543, 0.959281329, 0.954994135, 0.950256052, 0.945019662, 0.939232555, 0.932836813, 0.925768425, 0.917956648, 0.9093233, 0.899781974, 0.889237179, 0.877583378, 0.864703937, 0.850469953, 0.834738968, 0.817353541, 0.798139673, 0.776905065, 0.753437194, 0.727501186, 0.698837465, 0.667159154, 0.632149207, 0.593457232, 0.550695988, 0.503437504, 0.451208803, 0.393487162, 0.329694885, 0.259193516, 0.181277454, 0.09516689, 4.52654407e-18])
>>> assert np.allclose(bvp_solution, __bvp_solution, rtol=0.01, atol=0.01), 'bvp_solution is incorrect.'
```


Tests for q2_1 (points: None)
```
>>> assert np.isclose(du_dx_0, -0.001), f'`du_dx_0` should be -0.001, got {du_dx_0}.'
>>> assert np.isclose(u1, 0.0), f'`u1` should be 0.0, got {u1}.'
```
```
>>> assert isinstance(shooting_solution2, np.ndarray), '`shooting_solution` must be a NumPy array.'
>>> assert shooting_solution2.shape == x.shape, '`shooting_solution` must have the same length as `x`.'
>>> __shooting_solution2 = np.array([2.1988225, 2.19881199, 2.19880036, 2.19878752, 2.19877332, 2.19875763, 2.19874028, 2.1987211, 2.19869992, 2.19867654, 2.1986507, 2.19862216, 2.19859061, 2.19855572, 2.19851713, 2.19847445, 2.19842727, 2.19837513, 2.19831754, 2.19825382, 2.19818342, 2.19810571, 2.19801996, 2.19792531, 2.19782076, 2.19770522, 2.19757745, 2.1974361, 2.19727971, 2.19710667, 2.19691527, 2.19670367, 2.19646992, 2.19621183, 2.19592624, 2.1956107, 2.19526249, 2.19487834, 2.19445439, 2.19398626, 2.19346896, 2.19289697, 2.19226418, 2.19156394, 2.19078903, 2.18993165, 2.18898345, 2.18793553, 2.1867784, 2.18550002, 2.18408561, 2.1825238, 2.18080075, 2.17889989, 2.17680194, 2.17448488, 2.17192399, 2.16909182, 2.1659582, 2.16249024, 2.15865235, 2.15440618, 2.14971068, 2.1445221, 2.13879393, 2.13246268, 2.12545951, 2.1177282, 2.10919938, 2.09979037, 2.08940514, 2.07793438, 2.06525545, 2.05123237, 2.03571585, 2.01854331, 1.9995388, 1.97851309, 1.95526361, 1.92957448, 1.90121649, 1.86986502, 1.83519227, 1.79691794, 1.75469634, 1.70811628, 1.65670108, 1.59990862, 1.53713128, 1.46769597, 1.39086411, 1.30583167, 1.21172913, 1.1076215, 0.99250829, 0.865323566, 0.724935511, 0.56991247, 0.398588726, 0.20924783, -8.92024836e-06])
>>> assert np.allclose(shooting_solution2, __shooting_solution2, rtol=0.01, atol=0.01)
```
```
>>> assert isinstance(analytical_solution2, np.ndarray), '`analytical_solution` must be a NumPy array.'
>>> assert analytical_solution2.shape == x.shape, '`analytical_solution` must have the same length as `x`.'
>>> __analytical_solution2 = np.array([2.20254658, 2.20253606, 2.20252444, 2.20251159, 2.2024974, 2.20248171, 2.20246437, 2.2024452, 2.20242403, 2.20240062, 2.20237475, 2.20234616, 2.20231457, 2.20227965, 2.20224106, 2.20219841, 2.20215128, 2.20209918, 2.20204161, 2.20197799, 2.20190767, 2.20182996, 2.20174408, 2.20164916, 2.20154426, 2.20142833, 2.20130021, 2.20115861, 2.20100211, 2.20082916, 2.20063803, 2.20042678, 2.20019333, 2.19993532, 2.19965017, 2.19933503, 2.19898676, 2.19860185, 2.19817646, 2.19770633, 2.19718676, 2.19661255, 2.19597795, 2.1952766, 2.19450149, 2.19364487, 2.19269815, 2.19165186, 2.19049554, 2.1892176, 2.18780526, 2.18624439, 2.18451936, 2.1826129, 2.18050594, 2.17817739, 2.17560394, 2.17275984, 2.16961662, 2.16614283, 2.1623037, 2.1580608, 2.15337168, 2.14818939, 2.14246208, 2.13613242, 2.12913706, 2.121406, 2.11286185, 2.10341911, 2.09298326, 2.08144987, 2.0687035, 2.05461659, 2.03904814, 2.02184234, 2.00282699, 1.98181178, 1.95858638, 1.93291835, 1.90455078, 1.87319977, 1.83855155, 1.80025934, 1.7579399, 1.7111697, 1.65948062, 1.60235536, 1.53922218, 1.46944923, 1.39233819, 1.30711731, 1.21293367, 1.10884466, 0.99380851, 0.8666739, 0.72616842, 0.57088586, 0.39927209, 0.20960954, -0.0])
>>> assert np.allclose(analytical_solution2, __analytical_solution2, rtol=0.01, atol=0.01), ' `analytical_solution` solution does not match expected values.'
```


Tests for q2_2 (points: None)
```
>>> u_left_test = np.array([1.0, -0.001])
>>> u_right_test = np.array([0.0, 0.0])
>>> result = bc2(u_left_test, u_right_test)
>>> assert isinstance(result, np.ndarray), '`bc2` must return a NumPy array.'
>>> assert result.shape == (2,), f'`bc2` must return a 2-element array, got shape {result.shape}.'
>>> expected = np.array([u_left_test[1] + 0.001, u_right_test[0]])
>>> assert np.allclose(result, expected), f'`bc2` output is incorrect.'
```
```
>>> assert isinstance(bvp_solution2, np.ndarray), '`bvp_solution2` must be a NumPy array.'
>>> assert bvp_solution2.shape == x.shape, '`bvp_solution2` must have the same length as `x`.'
>>> __bvp_solution2 = np.array([2.20254352, 2.202533, 2.20252138, 2.20250853, 2.20249434, 2.20247865, 2.20246131, 2.20244214, 2.20242096, 2.20239756, 2.20237169, 2.2023431, 2.20231151, 2.20227659, 2.202238, 2.20219535, 2.20214822, 2.20209612, 2.20203855, 2.20197493, 2.20190461, 2.2018269, 2.20174102, 2.2016461, 2.2015412, 2.20142527, 2.20129715, 2.20115555, 2.20099905, 2.2008261, 2.20063497, 2.20042372, 2.20019027, 2.19993226, 2.19964711, 2.19933197, 2.1989837, 2.19859879, 2.1981734, 2.19770328, 2.19718371, 2.19660949, 2.19597489, 2.19527354, 2.19449844, 2.19364181, 2.19269509, 2.19164881, 2.19049248, 2.18921455, 2.18780221, 2.18624134, 2.18451631, 2.18260985, 2.18050289, 2.17817434, 2.1756009, 2.1727568, 2.16961359, 2.1661398, 2.16230067, 2.15805778, 2.15336866, 2.14818638, 2.14245907, 2.13612942, 2.12913407, 2.12140301, 2.11285887, 2.10341614, 2.09298031, 2.08144693, 2.06870058, 2.05461368, 2.03904524, 2.02183947, 2.00282414, 1.98180896, 1.95858359, 1.93291558, 1.90454805, 1.87319708, 1.8385489, 1.80025674, 1.75793736, 1.71116721, 1.65947821, 1.60235302, 1.53921993, 1.46944707, 1.39233614, 1.30711538, 1.21293188, 1.10884301, 0.993807025, 0.866672599, 0.726167332, 0.570884998, 0.399271482, 0.209609217, 4.11996826e-18])
>>> assert np.allclose(bvp_solution2, __bvp_solution2, rtol=0.01, atol=0.01), 'bvp_solution2 is incorrect.'
```


Tests for q3_1 (points: None)
```
>>> __u_fd = np.array([1.0, 0.99999526, 0.99999002, 0.99998423, 0.99997783, 0.99997076, 0.99996294, 0.9999543, 0.99994475, 0.9999342, 0.99992253, 0.99990964, 0.99989539, 0.99987964, 0.99986223, 0.99984299, 0.99982172, 0.99979821, 0.99977223, 0.99974352, 0.99971178, 0.9996767, 0.99963793, 0.99959508, 0.99954772, 0.99949537, 0.99943751, 0.99937356, 0.99930288, 0.99922476, 0.99913841, 0.99904298, 0.9989375, 0.99882092, 0.99869207, 0.99854965, 0.99839224, 0.99821827, 0.99802598, 0.99781345, 0.99757854, 0.99731891, 0.99703195, 0.99671479, 0.99636424, 0.99597679, 0.99554855, 0.99507524, 0.9945521, 0.9939739, 0.99333483, 0.9926285, 0.99184781, 0.99098495, 0.99003125, 0.98897717, 0.98781214, 0.98652446, 0.98510125, 0.98352822, 0.98178961, 0.97986798, 0.97774408, 0.97539662, 0.97280205, 0.96993436, 0.96676482, 0.96326164, 0.95938971, 0.9551102, 0.95038022, 0.94515234, 0.93937417, 0.93298776, 0.9259291, 0.91812743, 0.90950452, 0.89997394, 0.88944014, 0.87779753, 0.86492937, 0.85070667, 0.83498684, 0.81761229, 0.79840885, 0.77718399, 0.75372493, 0.7277965, 0.69913876, 0.66746442, 0.63245593, 0.59376234, 0.55099575, 0.5037274, 0.45148344, 0.39374012, 0.32991855, 0.25937892, 0.18141406, 0.09524238, 0.0])
>>> assert np.allclose(u_fd, __u_fd, rtol=0.01, atol=0.01), 'u_fd solution is wrong.'
```
