# CP 2024-25: Q2 Lecture 5 - Advanced Optimization (Part 2)

### General Guidelines

> ⚠️⚠️⚠️ READ CAREFULLY ⚠️⚠️⚠️

- Do not add, delete or create cells, write the answer only in the space marked with the three dots (`...`). Where function skeletons are provided, it is assumed that that function can be called again with different inputs somewhere else. So be careful to write code outside of functions.
  - Function should be ['pure'](https://en.wikipedia.org/wiki/Pure_function), thus no side effects, unless otherwise specified.
- Run the the first cell to import all libraries when opening the notebook before running your own code.
- Read carefully what is required to be printed/returned/plotted in the answer. Please do not output what is not asked for. 
  - If you used the print function for debugging, comment it out ( Ctlr + / ) before submitting
- All plots should have title, xlabel, ylabel, and legend (if there are more than one curve on the plot)
- Use the `help()` function, consult python documentation when using new functions, or do a web search and consult [stackoverflow](https://stackoverflow.com/questions/tagged/python)
- Please read the error messages if you get any, and try to understand what they mean. Debugging code is an essential skill to develop.
- You can use `%debug` to start an IPython console in a cell (or a scratchpad cell!) after an exception has occurred to try to debug.
- You can use `%pdb` to toggle the Python DeBugger (pdb) auto start after an unhandled exception.
- In the assignments you will find some tests put in place, to help you verify your solution. If these fail you are certain you did something wrong, thus look at the hints they provide. But passing these tests does __not__ mean your solution is actually correct.

Make sure you use `python3.11` and the package versions as stated in the provided `requirements.txt`. This file should also be on the course page.

```python
# Importing relevant libraries in the assignment

# This will create static plots (no zooming etc.)
# otherwise try just plain `%matplotlib`, or install a backend such as ipympl or PyQt5 and
# do or `%matplotlib ipympl` `%matplotlib qt`
%matplotlib inline

REPEAT_IMPORTS = True

if REPEAT_IMPORTS or ("IMPORTED_ALL" not in globals()):  # To save you a bit of time

    def print_import_info(package):
        print(
            "Successfully imported %-15s \tVersion: %10s"
            % (package.__name__, package.__version__)
        )

    ### Standard library imports

    import sys

    print("Python version {}".format(sys.version))
    if sys.version_info < (3, 8):
        print(
            "\u001b[31m"  # red
            "\u001b[1m"  # bold
            "WARNING: Use Python 3.8 or newer not to encounter any errors or problems later on. You can chance the the version. This sometime can be done by switching the kernel under the 'Kernel' tab."
            "\u001b[0m"  # reset
        )
    del sys  # Do not need it anymore

    # import functools
    import typing
    from typing import Dict, List, Tuple, Union

    ### Import third party libraries
    # Initialize self assessment helper
    import otter

    grader = otter.Notebook("Assignment_Q2_L5.ipynb")

    import os

    import numpy as np
    import numpy.typing as npt

    print_import_info(np)

    import scipy
    import scipy.optimize

    print_import_info(scipy)

    IMPORTED_ALL = True
    print("Finished importing packages")
else:
    print("Already imported all packages")
```

## Question 1 - Optimal blending problem.
Cattle feed can be mixed from oats, corn, alfalfa and peanut hulls.
The following table shows the current cost per ton (in dollars) of each of these ingredients, together with percentage of recommended daily allowances for protein fat, and fiber that a serving fulfills.


\begin{array}{|l|c|c|c|c|}
\hline
Amount (tons)& \textbf{Oats} & \textbf{Corn} & \textbf{Alfalfa} & \textbf{Hulls} \\ \hline
\% \text{Protein} & 60 & 80 & 55 & 40 \\ \hline
\% \text{Fat} & 50 & 70 & 40 & 100 \\ \hline
\% \text{Fiber} & 90 & 30 & 60 & 80 \\ \hline
\text{Cost} & 200 & 150 & 100 & 75 \\ \hline
\end{array}

We want to find a minimum cost way to produce 1 ton of feed that satisfies at least 60% of the daily allowance for protein and fiber while not exceeding 60 % of the fat allowance.

Formulate a linear program to choose an optimal feed mix that minimizes the cost of the blending mix while adhering to the above-mentioned specifications.


### Numerical solution strategies
Solve the formulated linear programming problem using **`scipy.optimize.linprog`**

### Question 1.1: Formulate the linear programming problem
Define a set $\mathcal{I}$ which comprises the set of ingredients.<br>
For each element of $\mathcal{I}$ define a variable $x$ that determines the amount of ingredient to be added to the feed mix.<br>
Define a parameter $c$ that stores the cost per unit of each ingredient. <br>
Formulate the objective such that it sums up the the product of the cost of an ingredient and the number of units of each ingredient.<br>
Formulate the composition constraint individually for protein, fat, and fiber according to the information provided in the table. <br>
For each nutrient, sum up over each ingredient and % of nutrient it provides such that minimum intake requirements for protein, fat, and fiber are met.<br>


Write down your formulation in this markdown cell below. <br>





#### Question 1.2 - Implement your formulation and solve it using SciPy
**Instructions**
You can also use the implementation shown during the live coding exercise as a guide on implementing the formulation for this exercise in SciPy.

1. **Create a numpy array to store the cost coefficients for the objective function**:
   - create a `numpy` array named **`cost`** to store the cost coefficients for oats, corn, alfalfa and hulls respectively.

2. **Consider to create two numpy arrays ($A$ and $b$) to express the inequality constraints in your formulation in the form $Ax \leq b$**:
   - Create a numpy array of dimensions: (number of nutrients $\times$ number of ingredients) named `A` to store the left-hand side coefficients of the inequality constraints
   - Create a numpy array of dimensions: (1 $\times$ number of nutrients) named `b` to store the right-hand side coefficients of the inequality constraints
   - In case you have a constraint which is of the type $\geq$ consider multiplying both left and right hand side by -1 to derive the equivalent form conforming to $Ax \leq b$

3. **Define two numpy arrays (`A_eq` and `b_eq`) to define the equality constraints in your formulation**

3. **Specify the bounds as a tuple or a list of tuples according to your formulation**
   - Store the bounds in a tuple/list of tuples named `bounds`

4. **Call `scipy.optimize.linprog`** to solve the formulation you derived and store the solution in a variable named `result`
   - Pass the cost coefficients, matrices defining the constraints and bounds to the linprog function.
   - You can refer to the documentation of `linprog` in SciPy for further help. You can access the documentation of **`scipy.optimize.linprog`** [here](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.linprog.html).

##### TIP: For the tests of the grader to pass you need to ensure that the order of the constraints in the matrices A and b is the same as implemented in the tests.
Please consider defining the inequality constraints in the following order:
1. Constraint to ensure minimum intake requirement of protein
2. Constraint to ensure maximum intake requirement of fat
3. Constraint to ensure minimum intake requirement of fiber

```python
# Import the necessary functions
"""
Cattle feed:

Objective:
    find ka minimum cost way to produce 1 ton of feed that satisfies at least 60% of the daily allowance for protein and fiber while not exceeding 60% of the fat allowance.

Solver used:
    scipy.optimize.linprog

Define a set L which comporises the set iof ingredients. 
    - For each element of L define a variable x that determines the amount of ingredient to be added to the feed mix.
    - Define a parameter c that stores the cost per unit of each ingredient
    - Formulate the objective such that it sums up the product of the cost of an ingredient and the number of units of each igredient.
    - Formulate the composition constraint individually for protein, fat and fiber according to the information provided in the table.
    Fro each nutient, sum up over each ingredient and % of nutrient it provides such that minimum intake requirements for protein, fat, and fiber are met.
    
"""

...
from scipy.optimize import linprog
```

```python
# Define the coefficients for the objective function
cost = np.ndarray([200,150,100,75])

```

```python
# Defining the inequality constraints in the form Ax <= b
A = np.ndarray([[60,50,90],
                [80,70,30],
                [55,40,60],
                [40,100,80]])

b = np.ndarray([0,0,0,0])


# since the inequality is Ax => b
# it needs to be changed to =<, this can be done by multiplying both sides by -1
A = A*-1
b = b*-1

```

```python
# Coefficients of the equality constraints (A_eq x = b_eq)
A_eq = ""
b_eq = ""
```

```python
# Providing the bounds on decision variables

# call scipy.optimize.linprog to solve the formulation you derived and store the solution in a varialbe named result
#   - pass coefficient matrixes defining the constraints and bounds to the linprog function. 
```

```python
# Calling the function `linprog` from scipy
# Pass all the arrays to the linprog function
result = linprog(
    cost, A_ub=A, b_ub=b, A_eq=A_eq, b_eq=b_eq, bounds=bounds, method="highs"
...
```

```python
# Printing the solution
# Perform the minimization
print("Cost of the blending mix ($):", result.fun)
print("Optimal solution:")
print("Amount of oats (tons):", result.x[0])
print("Amount of corn (tons):", result.x[1])
print("Amount of alfalfa (tons):", result.x[2])
print("Amount of hulls (tons):", result.x[3])
```

```python
result
```

**Reflect on the results**:<br>
1. Please provide a logical explanation of your optimal solution below:

2. Are any constraints active? How do you think the value of the objective will change on changing the bounds on the active constraint? 

```python
grader.check("q1")
```

### Question 2: Evaluating whether to build a carbon capture plant or not.
#### Statement
A carbon capture initiative is evaluating two potential sites, Site A and Site B, for the installation of carbon capture plants.<br> 
These plants are tasked with capturing carbon dioxide (CO₂) and delivering it to two locations, C and D where CO₂ electrolyzers have been built.<br> 
The minimum amount of CO₂ that can be electrolyzed at a given location and the associated costs of transporting CO₂ are detailed below:<br>
| Plant($/ton) | C | D |
|------|----------------------|----------------------|
| A    | 1.00                | 3.00                |
| B    | 4.50                | 1.00                |

The maximum capacity of both plants A and B to capture CO₂ is 500 tons per day if built. \
The existing CO₂ electrolyzer plants at sites C and D need to process a minimum of 200 and 250 tons of CO₂ per day, respectively for the process to be economical. \
Fixed costs for setting up plants at A and B are 700 and 610 $/day. \
We need to decide which plants to build and how much CO₂ they should capture with minimum total cost. <br>
Implement this statement as a mixed-integer linear program and solve it using SciPy.

Please consider the definitions below:<br>
##### Sets -
$i \in \mathcal{I}$ - two potential plant sites A and B <br>
$j \in \mathcal{J}$ - two CO₂ electrolyzer sites C and D <br>
##### Parameters -
$c_{ij}$ - unit cost of transporting CO₂ from plant $i$ to electrolyzer site $j$ <br>
$C_{i}$ - fixed cost of building plant $i$ <br>
$R_j$ - minimum processing capacity of electrolyzer $j$ <br>
$Q_i$ - capacity of plant $i$ to capture CO₂ <br>
##### Variables -
$y_i$ - decision variable (binary variable) associated with whether to build a plant A or B <br>
$S_{ij}$ - quantity of CO₂ captured at plant $i$ and sent to electrolyzer site $j$ <br>

#### Instructions
1. **Formulating the optimization problem**:
    - Consider formulating the problem using pen and paper first.
    - Use the techniques covered in the lecture to model logical constraints

2. **Implement the formulation in SciPy and then solve**
    - You can use the live-coding exercise as a guide on implementing mixed-integer linear programs in SciPy.
    - You need to import the following:<br>
      a. [milp](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.milp.html), <br>
      b. [Bounds](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.Bounds.html) and, <br>
      c. [LinearConstraint](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.LinearConstraint.html#scipy.optimize.LinearConstraint)

3. **Use the naming convention as in the previous exercise**
    - For defining the vector of binary variables define a list named `integrality`
    - Here the bounds are passed using Bounds method in `scipy.optimize`.

##### TIP: For the tests of the grader to pass you need to ensure that the order of the constraints in the matrices A and b is the same as implemented in the tests.
Please consider defining the inequality constraints in the following order:
1. Constraint to ensure minimum processing capacity of CO₂ electrolyzer at C
2. Constraint to ensure minimum processing capacity of CO₂ electrolyzer at D
3. Constraint to ensure maximum capacity of plant A is adhered to if built
4. Constraint to ensure maximum capacity of plant B is adhered to if built


```python
# Import the necessary functions
...
```

```python
# Coefficients for the objective function
...
```

```python
# Define the inequality constraints (Ax <= b)
...
```

```python
# Define the bounds for the variables
...
```

```python
# Define the integrality constraints
...
```

```python
# Calling the function `optimize_value_of_FA`
# Here you can adjust the bounds, initial guess and the solver as you wish by changing the values of the arguments passed to the function
result = milp(
    c=cost,
    constraints=LinearConstraint(A, -np.inf, b),
    bounds=bounds,
    integrality=integrality,
...
```

```python
# Print the results
# Check the result and print the output
if result.success:
    print("Status:", result.message)
    print("Objective value:", result.fun)
    print("Variable values:", result.x)
else:
    print("Optimization failed.")
    print("Status:", result.message)
```

**Reflect on the results**:<br>
1. Please provide a logical explanation of your optimal solution below:

2. Are any constraints active? How do you think the value of the objective will change on changing the bounds on the active constraint? 

```python
grader.check("q2")
```

Tests for q1 (points: 0)
```
>>> expected_cost = np.array([200, 150, 100, 75])
>>> sorted_cost = np.sort(cost)[::-1]
>>> assert np.allclose(expected_cost, sorted_cost, atol=1e-10), f'Cost coefficients not defined correctly..'
```
```
>>> expected_A_ub = np.array([[-0.4, -0.55, -0.6, -0.8], [1, 0.7, 0.5, 0.4], [-0.3, -0.6, -0.8, -0.9]])
>>> sorted_A_ub = np.sort(A, axis=1)[:, ::-1]
>>> assert np.allclose(expected_A_ub, sorted_A_ub, atol=1e-10), f'Left-hand side of inequality constraints not defined correctly.'
```
```
>>> expected_b_ub = np.array([0.6, -0.6, -0.6])
>>> sorted_b_ub = np.sort(b)[::-1]
>>> assert np.allclose(expected_b_ub, sorted_b_ub, atol=1e-10), f'Right-hand side of inequality constraints not defined correctly.'
```
```
>>> expected_A_eq = np.array([[1, 1, 1, 1]])
>>> assert np.allclose(expected_A_eq, A_eq, atol=1e-10), f'Left-hand side of equality constraints not defined correctly.'
```
```
>>> expected_b_eq = np.array([1])
>>> assert np.allclose(expected_b_eq, b_eq, atol=1e-10), f'Left-hand side of equality constraints not defined correctly.'
```
```
>>> expected_optimal_cost = 125
>>> output = result.fun
>>> assert np.allclose(expected_optimal_cost, output, atol=0.001), f'For the optimization call, the solution is not within tolerance. Make sure you have defined the constraints correctly.'
```


Tests for q2 (points: 0)
```
>>> expected_cost = np.array([700, 610, 4.5, 3, 1, 1])
>>> sorted_cost = np.sort(cost)[::-1]
>>> assert np.allclose(expected_cost, sorted_cost, atol=1e-10), f'Cost coefficients not defined correctly..'
```
```
>>> expected_A_ub = np.array([[0, 0, 0, 0, -1, -1], [0, 0, 0, 0, -1, -1], [1, 1, 0, 0, 0, -500], [1, 1, 0, 0, 0, -500]])
>>> sorted_A_ub = np.sort(A, axis=1)[:, ::-1]
>>> assert np.allclose(expected_A_ub, sorted_A_ub, atol=1e-10), f'Left-hand side of inequality constraints not defined correctly.'
```
```
>>> expected_b_ub = np.array([0, 0, -200, -250])
>>> sorted_b_ub = np.sort(b)[::-1]
>>> assert np.allclose(expected_b_ub, sorted_b_ub, atol=1e-10), f'Right-hand side of inequality constraints not defined correctly.'
```
```
>>> expected_integrality = np.array([1, 1, 0, 0, 0, 0])
>>> sorted_integrality = np.sort(integrality)[::-1]
>>> assert np.allclose(expected_integrality, sorted_integrality, atol=1e-10), f'Integrality of variables not defined correctly.'
```
```
>>> expected_optimal_cost = 1650
>>> output = result.fun
>>> assert np.allclose(expected_optimal_cost, output, atol=0.001), f'For the optimization call, the solution is not within tolerance. Make sure you have defined the constraints correctly.'
```
